# Analiza danych przestrzennych - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają na celu wyłapać podstawowe błędy. Przejście przez kod wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem). Rozmiar poziomy figur nie powinien przekraczać 20 cali.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł.
- Zakresy osi wykresów przedstawiających rozkłady punktów w przestrzeni powinny być docięte do granic obszaru, na którym został wygenerowany proces punktowy.
- Proporcje osi wykresów przedstawiających rozkłady punktów w przestrzeni powinny być dobrane tak, aby wykresy odzwierciedlały rzeczywisty rozkład punktów w przestrzeni. Nie należy osiągać tego efektu manipulując rozmiarem całej figury.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu:
Członkowie:

---

## Zestaw zadań 1: Generowanie procesów punktowych

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Proces regularny z komponentem losowym w obrębie prostokątnego obszaru [7 pkt]

#### a) Przygotowanie funkcji
Przygotuj funkcję `regular_on_rectangle()`, która będzie generować proces regularny z komponentem losowym w obrębie prostokątnego obszaru.

Algorytm postępowania:
1) Generujemy punkty na regularnej siatce tak, że poszczególne punkty są oddalone od siebie o odległość $d_x$ w poziomie i $d_y$ w pionie, a skrajne punkty są oddalone od krawędzi obszaru odpowiednio o $\frac{1}{2} d_x$ w poziomie i $\frac{1}{2}d_y$ w pionie.
2) Do współrzędnych każdego z punktów dodajemy komponent losowy określony wzorem $r \cdot rand$ , gdzie $r$ to parametr określający wielkość komponentu losowego, a $rand$ to wartość wylosowana z rozkładu jednostajnego na przedziale $[-\frac{1}{2} d_x, \frac{1}{2} d_x]$ w przypadku współrzędnej $x$ i na przedziale $[-\frac{1}{2} d_y, \frac{1}{2} d_y]$ w przypadku współrzędnej $y$.

<!-- grade_id: cell-0b8ca172a3462f06, points: 0, grade: false, solution: true -->
```python
def regular_on_rectangle(grid, random_component, x_lim, y_lim):        
    """
    Parameters
    -------
    grid: list
        Lista określająca liczbę punktów w siatce w poziomie i w pionie.
        Przykład: [20, 10]
    random_component: float
        Liczba z przedziału [0, 1] określająca wielkość komponentu losowego.
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10]   

    Returns
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    """
    # calculate distance between points
    dx = ( x_lim[1] - x_lim[0] ) / grid[0]
    dy = ( y_lim[1] - y_lim[0] ) / grid[1]
    
    # calculate all values of x and y coordinates
    x_coords = np.linspace( x_lim[0] + dx, x_lim[1] - dx, grid[0] )
    y_coords = np.linspace( y_lim[0] + dy, y_lim[1] - dy, grid[1] )

    # calculate random component for each point's x and y value
    x_matrix, y_matrix = np.meshgrid( x_coords, y_coords )
    x_noise = random_component * np.random.uniform( -1/2*dx, 1/2*dx, grid[0] * grid[1] ) 
    y_noise = random_component * np.random.uniform( -1/2*dy, 1/2*dy, grid[0] * grid[1] )

    # add noise (random component) to x and y coordinates
    x_list = x_matrix.flatten()
    y_list = y_matrix.flatten()
    x_list += x_noise
    y_list += y_noise
    
    # create dataframe out of coordinates
    points = pd.DataFrame( { "X" : x_list, "Y" : y_list } )

    return points
```

<!-- grade_id: cell-39c0feb5c4d29909, points: 5, grade: true, solution: false -->
```python
# Komórka testowa
test_points_1 = regular_on_rectangle(grid=[30, 15], random_component=0.25, x_lim=[-5, 15], y_lim=[20, 25])
assert type(test_points_1) == pd.DataFrame
assert test_points_1.shape == (450, 2)
assert list(test_points_1.columns) == ["X", "Y"]
assert np.min(test_points_1["X"]) >= -5
assert np.max(test_points_1["X"]) <= 15
assert np.min(test_points_1["Y"]) >= 20
assert np.max(test_points_1["Y"]) <= 25

test_points_2 = regular_on_rectangle(grid=[50, 50], random_component=0.47, x_lim=[-10, 10], y_lim=[-30, -10])
assert type(test_points_2) == pd.DataFrame
assert test_points_2.shape == (2500, 2)
assert list(test_points_2.columns) == ["X", "Y"]
assert np.min(test_points_2["X"]) >= -10
assert np.max(test_points_2["X"]) <= 10
assert np.min(test_points_2["Y"]) >= -30
assert np.max(test_points_2["Y"]) <= -10
```

#### b) Wygenerowanie procesu i wizualizacja

Wykorzystaj przygotowaną funkcję do wygenerowania następującego rozkładu punktów i przedstaw go w postaci wykresu punktowego.

Parametry rozkładu:
 - liczba punktów w poziomie: 20
 - liczba punktów w pionie: 10
 - wartość współczynnika określającego siłę komponentu losowego: 0.5
 - zakres wartości współrzędnej x: $[20, 40]$
 - zakres wartości współrzędnej y: $[-10, 0]$

<!-- grade_id: cell-cdcecd5de1baaf3b, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
plt.subplots(figsize=(16,6))
points = regular_on_rectangle(grid=[20, 10], random_component=0.5, x_lim=[20, 40], y_lim=[-10, 0])
plt.plot(points["X"], points["Y"], marker='o', color='k', linestyle='none')
plt.title("Proces regularny z komponentem losowym ",fontsize=16)
plt.xlabel("X",fontsize=14)
plt.ylabel("Y",fontsize=14)
plt.xlim([20,40])
plt.ylim([-10,0])
plt.show()
```

### Zadanie 2: Jednorodny proces punktowy Poissona w obrębie prostokątnego obszaru [7 pkt]

#### a) Przygotowanie funkcji
Przygotuj funkcję `homogeneous_poisson_on_rectangle()`, która będzie generować jednorodny proces punktowy Poissona w obrębie prostokątnego obszaru.

Algorytm postępowania:
1. Losujemy liczbę punktów $n$ z rozkładu Poissona o wartości oczekiwanej równej $i \cdot a$, gdzie $i$ to intensywność procesu, a $a$ jest powierzchnią obszaru, na którym generujemy punkty.
2. Niezależnie od siebie losujemy wartości współrzędnych $n$ punktów z rozkładów jednostajnych na przedziałach określonych przez przyjęte zakresy wartości współrzędnych $x$ i $y$.

<!-- grade_id: cell-496fc81b5caba229, points: 0, grade: false, solution: true -->
```python
def homogeneous_poisson_on_rectangle(intensity, x_lim, y_lim):
    """
    # liczba punktów intensywnosc-srednia oczekiwana wartosc w danym kwadraciku 
    # intensywnosc* Pole powierzchni - argument do funkcji z ktorej wyciagamy wartoscliczby puntkow  z rozkladu poissona 
    ich lokalizacja jest z rozkladu normalnego 
    Parameters
    
    im nisza intensywnosc funkcji intensywnosci 1-i/imax  (0,1) 
    R(x,y)=x
    ------- 
    intensity: float
        Liczba dodatnia określająca intensywność procesu punktowego.
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10]   
    
    Returns
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    """

    # calculate area to generate points on
    area = ( x_lim[1] - x_lim[0] ) * ( y_lim[1] - y_lim[0] )

    # draw number of points from Poisson distribution
    n_points = np.random.poisson( area * intensity, 1 )

    # draw points coordinates from uniform distribution
    x_points = np.random.uniform( x_lim[0], x_lim[1], n_points ) 
    y_points = np.random.uniform( y_lim[0], y_lim[1], n_points )

    # create dataframe out of coordinates
    points = pd.DataFrame( { "X" : x_points, "Y" : y_points } )

    return points
```

<!-- grade_id: cell-bd01fb3f7c8e0e46, points: 5, grade: true, solution: false -->
```python
# Komórka testowa
test_points_3 = homogeneous_poisson_on_rectangle(intensity=5, x_lim=[-11, 18], y_lim=[5, 10])
assert type(test_points_3) == pd.DataFrame
assert test_points_3.shape[1] == 2
assert list(test_points_3.columns) == ["X", "Y"]
assert np.min(test_points_3["X"]) >= -11
assert np.max(test_points_3["X"]) <= 18
assert np.min(test_points_3["Y"]) >= 5
assert np.max(test_points_3["Y"]) <= 10

test_points_4 = homogeneous_poisson_on_rectangle(intensity=2.5, x_lim=[13 ,27], y_lim=[-15, -8])
assert type(test_points_4) == pd.DataFrame
assert test_points_4.shape[1] == 2
assert list(test_points_4.columns) == ["X", "Y"]
assert np.min(test_points_4["X"]) >= 13
assert np.max(test_points_4["X"]) <= 27
assert np.min(test_points_4["Y"]) >= -15
assert np.max(test_points_4["Y"]) <= -8
```

#### b) Wygenerowanie procesu i wizualizacja
Wykorzystaj przygotowaną funkcję do wygenerowania następującego rozkładu punktów i przedstaw go w postaci wykresu punktowego.

Parametry rozkładu:
 - intensywność procesu: 10
 - zakres wartości współrzędnej x: $[-10, 10]$
 - zakres wartości współrzędnej y: $[5, 15]$

<!-- grade_id: cell-720584fea7bd395c, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
plt.subplots(figsize=(16,8))
points = homogeneous_poisson_on_rectangle(intensity=10, x_lim=[-10, 10], y_lim=[5, 15])
plt.plot(points["X"],points["Y"],marker='o',linestyle='none')
plt.title("Jednorodny proces punktowy Poissona ",fontsize=16)
plt.xlabel("X",fontsize=14)
plt.ylabel("Y",fontsize=14)
plt.xlim([-10,10])
plt.ylim([5,15])
plt.show()
```

### Zadanie 3: Niejednorodny proces punktowy Poissona w obrębie prostokątnego obszaru [12 pkt]

#### a) Przygotowanie funkcji
Przygotuj funkcję `unhomogeneous_poisson_on_rectangle()`, która będzie generować niejednorodny proces punktowy Poissona w obrębie prostokątnego obszaru.

Algorytm postępowania:
1. Generujemy punkty z wykorzystaniem jednorodnego procesu Poissona przyjmując intensywność procesu równą maksymalnej wartości jaką osiąga funkcja intensywności w analizowanym obszarze ($i_{max}$).
2. Przerzedzamy rozkład usuwając z niego punkty z prawdopodobieństwem równym $1-\frac{i}{i_{max}}$, gdzie $i$ to wartość jaką przyjmuje funkcja intensywności dla współrzędnych danego punktu, a $i_{max}$ to maksymalna wartość jaką funkcja intensywności osiąga w granicach analizowanego obszaru.


Uwagi do wykonania zadania:
- W celu znalezienia maksymanej wartości funkcji intensywności w obrębie obszaru należy skorzystać z funkcji `scipy.optimize.minimize()`.
- W części zadania dotyczącej generowania jednorodnego procesu punktowego Poissona nie należy dublować kodu napisanego w ramach zadania 2, a jedynie wywołać przygotowaną w ramach tego zadania funkcję `homogeneous_poisson_on_rectangle()`.

<!-- grade_id: cell-4258d05645de2b5e, points: 0, grade: false, solution: true -->
```python
def unhomogeneous_poisson_on_rectangle(intensity_function, x_lim, y_lim):
    """
    Parameters
    -------
    intensity_function: function
        Funkcja przyjmująca dwa argumenty (wartości współrzędnyych X i Y pojedynczego punktu
        lub dwie macierze 1D ze współrzędnymi X i Y wielu punktów) i zwracająca pojedynczą
        wartość lub macierz 1D wartości funkcji opisującej intensywność procesu dla tych współrzędnych.
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10] 
    
    Returns
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c637432eb96e6978, points: 10, grade: true, solution: false -->
```python
# Komórka testowa
def test_intensity_function(x, y):
    return x + 20

test_points_5 = unhomogeneous_poisson_on_rectangle(test_intensity_function, x_lim=[-15, 10], y_lim=[5, 15])
assert type(test_points_5) == pd.DataFrame
assert test_points_5.shape[1] == 2
assert list(test_points_5.columns) == ["X", "Y"]
assert np.min(test_points_5["X"]) >= -15
assert np.max(test_points_5["X"]) <= 10
assert np.min(test_points_5["Y"]) >= 5
assert np.max(test_points_5["Y"]) <= 15

test_points_6 = unhomogeneous_poisson_on_rectangle(test_intensity_function, x_lim=[7, 21], y_lim=[-14, -1])
assert type(test_points_6) == pd.DataFrame
assert test_points_6.shape[1] == 2
assert list(test_points_6.columns) == ["X", "Y"]
assert np.min(test_points_6["X"]) >= 7
assert np.max(test_points_6["X"]) <= 21
assert np.min(test_points_6["Y"]) >= -14
assert np.max(test_points_6["Y"]) <= -1
```

#### b) Wygenerowanie procesu i wizualizacja
Wykorzystaj przygotowaną funkcję do wygenerowania następującego rozkładu punktów i przedstaw go w postaci wykresu punktowego.

Parametry rozkładu:
 - intensywność procesu: f(x, y) = 2*x + 5*y
 - zakres wartości współrzędnej x: $[0, 20]$
 - zakres wartości współrzędnej y: $[0, 10]$

<!-- grade_id: cell-3b2e16df836b7a03, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Zadanie 4: Proces punktowy Matérna w obrębie prostokątnego obszaru [12 pkt]

#### a) Przygotowanie funkcji
Przygotuj funkcję `materna_on_rectangle()`, która będzie generować proces punktowy Materna w obrębie prostokątnego obszaru.

Algorytm postępowania:
1. Rozszerzamy analizowany obszar o bufor równy wartości parametru `cluster_radius`.
2. W obrębie rozszerzonego obszaru generujemy punkty w wykorzystaniem jednorodnego procesu Poissona o intensywności równej wartości parametru `parent_intensity`.
3. Wokół każdego z wygenerowanych punktów, w okręgu o promieniu równym wartości parametru `cluster_radius`, generujemy punkty z wykorzystaniem jednorodnego procesu Poissona o intensywności równej wartości parametru `daughter_intensity`.
4. Usuwamy z rozkładu wszystkie punkty leżące poza analizowanym (nierozszerzonym) obszarem.

Uwagi do wykonania zadania:
- Punkty wygenerowane w ramach procesu macierzystego (centra klastrów) nie są częścią wynikowego rozkładu punktów.
- W części zadania dotyczącej generowania jednorodnego procesu punktowego Poissona nie należy dublować kodu napisanego w ramach zadania 2, a jedynie wywołać przygotowaną w ramach tego zadania funkcję `homogeneous_poisson_on_rectangle()`.

<!-- grade_id: cell-37508c39f24a0801, points: 0, grade: false, solution: true -->
```python
def materna_on_rectangle(parent_intensity, daughter_intensity, cluster_radius, x_lim, y_lim):
    """
    Parameters
    -------
    parent_intensity: float
        Liczba dodatnia określająca intensywność macierzystego procesu punktowego.
    daughter_intensity: float
        Liczba dodatnia określająca intensywność potomnego procesu punktowego.
    cluster_radius: float
        Liczba dodatnia określająca promień generowanych klastrów.
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10]   
    
    rozkład zklasteryzowany; rozszerzamy obszar o l ktory ma generowac punkty o 
    w okol kazdego z tych punktow generujemy punkty z rozdkladu poisona o promieniu l  - mozna skorzystac z for; potem usuwamy ten nad obszar 
    *jak wygenerowac punkty w tym obszarze - wywolac funkcje z 2 zadania i przefiltrowac i usunac te ktore sa dalej niz r
    
    Returns
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    """
    points = homogeneous_poisson_on_rectangle(parent_intensity,[ x_lim[0] - cluster_radius, x_lim[1] + cluster_radius], [ y_lim[0] - cluster_radius, y_lim[1] + cluster_radius ])
    xy=pd.DataFrame({"X":[],"Y":[]})
    
    for i in range(len(points)):
        pp = homogeneous_poisson_on_rectangle(daughter_intensity,[points["X"][i]-cluster_radius,points["X"][i]+cluster_radius],[points["Y"][i]-cluster_radius,points["Y"][i]+cluster_radius])
        pp= pp[ np.sqrt( (pp["X"] - points["X"][i])** 2 + (pp["Y"] -  points["Y"][i])** 2 ) <= cluster_radius ]
        xy = pd.concat([xy, pp], ignore_index=True)
    return xy[ (xy["X"] >= x_lim[0]) & (xy["X"] <= x_lim[1]) & (xy["Y"] >= y_lim[0]) & (xy["Y"] <= y_lim[1]) ]
```

<!-- grade_id: cell-5000d87087be8234, points: 10, grade: true, solution: false -->
```python
# Komórka testowa
test_points_7 = materna_on_rectangle(parent_intensity=0.12, daughter_intensity=5, cluster_radius=1.2, x_lim=[-12, 17], y_lim=[-20,-4])
assert type(test_points_7) == pd.DataFrame
assert test_points_7.shape[1] == 2
assert list(test_points_7.columns) == ["X", "Y"]
assert np.min(test_points_7["X"]) >= -12
assert np.max(test_points_7["X"]) <= 17
assert np.min(test_points_7["Y"]) >= -20
assert np.max(test_points_7["Y"]) <= -4

test_points_8 = materna_on_rectangle(parent_intensity=0.4, daughter_intensity=20, cluster_radius=0.5, x_lim=[-25, -8], y_lim=[-2, 14])
assert type(test_points_8) == pd.DataFrame
assert test_points_8.shape[1] == 2
assert list(test_points_8.columns) == ["X", "Y"]
assert np.min(test_points_8["X"]) >= -25
assert np.max(test_points_8["X"]) <= -8
assert np.min(test_points_8["Y"]) >= -2
assert np.max(test_points_8["Y"]) <= 14
```

#### b) Wygenerowanie procesu i wizualizacja
Wykorzystaj przygotowaną funkcję do wygenerowania następującego rozkładu punktów i przedstaw go w postaci wykresu punktowego.

Parametry  rozkładu:
 - intensywność procesu macierzystego: 0.2
 - intensywność procesu potomnego: 10
 - promień klastra: 1
 - zakres wartości współrzędnej x: $[-15, 5]$
 - zakres wartości współrzędnej y: $[10, 20]$

<!-- grade_id: cell-7252b987abdc121f, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
plt.subplots(figsize=(16,8))
points = materna_on_rectangle(parent_intensity=0.2, daughter_intensity=10, cluster_radius=1, x_lim=[-15, 5], y_lim=[10,20])
plt.plot(points["X"],points["Y"],marker='o',linestyle='none')
plt.title("Proces punktowy Matérna",fontsize=16)
plt.xlabel("X",fontsize=14)
plt.ylabel("Y",fontsize=14)
plt.xlim([-15,5])
plt.ylim([10,20])
plt.show()
```

### Zadanie 5: Proces punktowy Thomasa w obrębie prostokątnego obszaru [12 pkt]

#### a) Przygotowanie funkcji
Przygotuj funkcję `thomas_on_rectangle()`, która będzie generować proces punktowy Thomasa w obrębie prostokątnego obszaru.

Algorytm postępowania:
1. Rozszerzamy analizowany obszar o bufor, którego wielkość jest zależna od wartości parametru `cluster_sigma` (na potrzeby zadania przyjmijmy 4-krotność wartości tego parametru).
2. W obrębie rozszerzonego obszaru generujemy punkty w wykorzystaniem jednorodnego procesu Poissona o intensywności równej wartości parametru `parent_intensity`.
3. Dla każdego z wygenerowanych punktów (centrów klastrów):
    - Losujemy liczbę punktów znajdujących się w obrębie danego klastra $n_c$ z rozkładu Poissona o wartości oczekiwanej równej wartości parametru `mean_cluster_size`.
    - Niezależnie od siebie losujemy wartości współrzędnych $n_c$ punktów z rozkładów normalnych o wartości oczekiwanej równej współrzędnej punktu stanowiącego centrum klastra i odchyleniu standardowym równym wartości parametru `cluster_sigma`.
4. Usuwamy z rozkładu wszystkie punkty leżące poza analizowanym (nierozszerzonym) obszarem.

Uwagi do wykonania zadania:
- Punkty wygenerowane w ramach procesu macierzystego (centra klastrów) nie są częścią wynikowego rozkładu punktów.
- W części zadania dotyczącej generowania jednorodnego procesu punktowego Poissona nie należy dublować kodu napisanego w ramach zadania 2, a jedynie wywołać przygotowaną w ramach tego zadania funkcję `homogeneous_poisson_on_rectangle()`.

<!-- grade_id: cell-4486abd11db366a8, points: 0, grade: false, solution: true -->
```python
def thomas_on_rectangle(parent_intensity, mean_cluster_size, cluster_sigma, x_lim, y_lim):
    """
    Parameters
    -------
    parent_intensity: float
        Liczba dodatnia określająca intensywność macierzystego procesu punktowego.
    mean_cluster_size: float
        Liczba dodatnia określająca oczekiwaną liczebność generowanych klastrów.
    cluster_sigma: float
        Liczba dodatnia określająca odchylenie standardowe rozkładu wykorzystywanego w procesie generowania klastrów.
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10]
        
    Returns
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    """
    # generate centers of clusters in broadened area
    centers = homogeneous_poisson_on_rectangle( parent_intensity, [x_lim[0] - 4 * cluster_sigma, x_lim[1] + 4 * cluster_sigma],
                                               [y_lim[0] - 4 * cluster_sigma, y_lim[1] + 4 * cluster_sigma] )
    
    # create empty dataframe to store points coords
    points = pd.DataFrame( columns= ["X", "Y"] )

    # draw points for each cluster center from normal distribution
    for center in centers.itertuples():
        # draw number of points for cluster from Poisson distribution
        n_points = np.random.poisson(mean_cluster_size)

        # draw points coord from normal distribution
        x_points = np.random.normal( center.X, cluster_sigma, n_points )
        y_points = np.random.normal( center.Y, cluster_sigma, n_points )
        
        # append to result dataframe
        xy_df = pd.DataFrame( { "X" : x_points, "Y" : y_points } )
        points = pd.concat( [xy_df, points] )
    
    # narrow area to original size
    points = points[ (points["X"] >= x_lim[0]) & (points["X"] <= x_lim[1]) & (points["Y"] >= y_lim[0]) & (points["Y"] <= y_lim[1]) ]
    return points
```

<!-- grade_id: cell-797f40b9afe7071f, points: 10, grade: true, solution: false -->
```python
# Komórka testowa
test_points_9 = thomas_on_rectangle(parent_intensity=0.13, mean_cluster_size=42, cluster_sigma=0.25, x_lim=[-12, 17], y_lim=[-20,-4])
assert type(test_points_9) == pd.DataFrame
assert test_points_9.shape[1] == 2
assert list(test_points_9.columns) == ["X", "Y"]
assert np.min(test_points_9["X"]) >= -12
assert np.max(test_points_9["X"]) <= 17
assert np.min(test_points_9["Y"]) >= -20
assert np.max(test_points_9["Y"]) <= -4

test_points_10 = thomas_on_rectangle(parent_intensity=0.54, mean_cluster_size=15, cluster_sigma=0.75, x_lim=[-25, -8], y_lim=[-2, 14])
assert type(test_points_10) == pd.DataFrame
assert test_points_10.shape[1] == 2
assert list(test_points_10.columns) == ["X", "Y"]
assert np.min(test_points_10["X"]) >= -25
assert np.max(test_points_10["X"]) <= -8
assert np.min(test_points_10["Y"]) >= -2
assert np.max(test_points_10["Y"]) <= 14
```

#### b) Wygenerowanie procesu i wizualizacja
Wykorzystaj przygotowaną funkcję do wygenerowania następującego rozkładu punktów i przedstaw go w postaci wykresu punktowego.

Parametry rozkładu:
 - intensywność procesu macierzystego: 0.2
 - średnia liczebność klastra: 30
 - odchylenie standardowe klastra: 0.5
 - zakres wartości współrzędnej x: $[25, 45]$
 - zakres wartości współrzędnej y: $[-10, 0]$

<!-- grade_id: cell-0e700b0a0fdef851, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
plt.subplots(figsize=(16,8))
points = thomas_on_rectangle(parent_intensity=0.2, mean_cluster_size=30, cluster_sigma=0.5, x_lim=[25, 45], y_lim=[-10,0])
plt.plot(points["X"],points["Y"],marker='o',linestyle='none')
plt.title("Proces punktowy Thomasa",fontsize=16)
plt.xlabel("X",fontsize=14)
plt.ylabel("Y",fontsize=14)
plt.xlim([25,45])
plt.ylim([-10,0])
plt.show()
```