# Analiza danych przestrzennych - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają na celu wyłapać podstawowe błędy. Przejście przez kod wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem). Rozmiar poziomy figur nie powinien przekraczać 20 cali.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł.
- Zakresy osi wykresów przedstawiających rozkłady punktów w przestrzeni powinny być docięte do granic obszaru, na którym został wygenerowany proces punktowy.
- Proporcje osi wykresów przedstawiających rozkłady punktów w przestrzeni powinny być dobrane tak, aby wykresy odzwierciedlały rzeczywisty rozkład punktów w przestrzeni. Nie należy osiągać tego efektu manipulując rozmiarem całej figury.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu: IAD22
Członkowie: Natalia Kubańska 411933, Nadia Napruszewska 411381, Kamila Żur 409777

---

## Zestaw zadań 2: Badanie intensywności procesów punktowych (część 1)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Wczytanie danych

Załączone do notatniki pliki zawierają następujące procesy punktowe:
 - `points_HP` - jednorodny rozkład Poissona,
 - `points_UP` - niejednorodny rozkład Poissona,
 - `points_M` - rozkład Materna,
 - `points_T` - rozkład Thomasa,
 
wygenerowane dla zakresu współrzędnych X $[-10, 10]$ i Y $[-5, 5]$.

W celu wczytania danych do notatnika umieść wszystkie pliki w tym samym folderze, w którym znajduje się notatnik.

```python
# Wczytanie danych
points_HP = pd.read_pickle('points_HP.pkl')
points_UP = pd.read_pickle('points_UP.pkl')
points_M = pd.read_pickle('points_M.pkl')
points_T = pd.read_pickle('points_T.pkl')
```

### Zadanie 1: Badanie intensywności procesu punktowego metodą szacowania lokalnego (15 pkt)

Przygotuj funkcję `point_count_on_subregions()`, która będzie zliczać punkty w obrębie prostokątnych podobszarów oraz funkcję `intensity_on_subregions()`, która będzie obliczać intensywność procesu punktowego w obrębie prostokątnych podobszarów.

Algorytm postępowania:
1. Dzielimy analizowany obszar na identyczne prostokątne podobszary ułożone w równomiernej siatce (w wierszach i kolumnach).
2. Zliczamy punkty znajdujące się w poszczególych podobszarach (punkty znajdujące się na granicach podobszarów powinny być zliczone tylko do jednego z nich).
3. Obliczamy intensywność procesu w obrębie poszczególnych podobszarów dzieląc znajdującą się w nich liczbę punktów przez pole powierzchni podobszaru.

Uwagi do wykonania zadania:
- W części zadania dotyczącej funkcji obliczającej intensywność procesu nie należy dublować kodu napisanego w funkcji zliczającej punkty w obrębie prostokątnych podobszarów, a jedynie wywołać przygotowaną wcześniej funkcję `point_count_on_subregions()`.
- W metodzie szacowania lokalnego szacowanie intensywności odbywa się na kwadratowych podobszarach. W celu ułatwienia pisania funkcji warunek ten w zadaniu został rozluźniony do obszarów prostokątnych. We wszystkich zadaniach rozmiar obszaru i parametry siatki podobszarów są natomiast dobrane w taki sposób, że wygenerowane podobszary będą kwadratami.

#### a) Przygotowanie funkcji

<!-- grade_id: cell-6991ddf3ffdd6af4, points: 0, grade: false, solution: true -->
```python
def point_count_on_subregions(points, bins, x_lim, y_lim):
    """
    Parameters
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    bins: list
        Lista określająca liczbę podobszarów w poziomie i pionie.
        Przykład: [10, 10]
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10]   

    Returns
    -------
    bin_data: list
        Lista zawierająca trzy macierze:
        - 2D z liczbą punków przypisanych do każdego z podobszarów.
        - 1D ze współrzędnymi krawędzi podobszarów na osi X,
        - 1D ze współrzędnymi krawędzi podobszarów na osi Y,
        Na przykład: [array([[7, 2], [4, 5]]), array([0, 1, 2]), array([0, 1, 2])]
            '''
    
    dzielimy obszar na kwadrty z punktami
    zliczamy ile jest punktow w danym kwadracie
    intensywnosc = ilosc punktow/ pole powierzchni 
    zakladamy ze moge byc prostokaty, ale dane sa dobrane tak ze wyjda nam kwadraty 
    pierwsza funkcja dzieli i zlicza; druga funkcja zlicza intensywnosc i wywolujemy w niej te pierwsza   
    '''
    """
    
    # YOUR CODE HERE
    H, xedges, yedges = np.histogram2d( points["X"], points["Y"], range = [ x_lim, y_lim ], bins = bins )
    
    return [ H.T, xedges, yedges ]
    

    

def intensity_on_subregions(points, bins, x_lim, y_lim):
    """
    Parameters
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    bins: list
        Lista określająca liczbę podobszarów w poziomie i pionie.
        Przykład: [10, 10]
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10]   

    Returns
    -------
    intensity_data: list
        Lista zawierająca trzy macierze:
        - 2D z wartością intensywności przypisą do każdego z podobszarów.
        - 1D ze współrzędnymi krawędzi podobszarów na osi X,
        - 1D ze współrzędnymi krawędzi podobszarów na osi Y,
        Na przykład: [array([[3, 7], [9, 4]]), array([0, 1, 2]), array([0, 1, 2])]
    """
    # YOUR CODE HERE
    H, xedges, yedges = point_count_on_subregions( points, bins, x_lim, y_lim )
    intensity = H / ( xedges[1]- xedges[0] ) * ( yedges[1] - yedges[0] )   
    
    return [ intensity, xedges, yedges ]

```

<!-- grade_id: cell-c751f7154b39b00e, points: 8, grade: true, solution: false -->
```python
# Komórka testowa
test_intensity_1 = intensity_on_subregions(points_HP, bins=[10, 5], x_lim=[-5, 5], y_lim=[-2.5, 2.5])
assert type(test_intensity_1) == list
assert len(test_intensity_1) == 3
assert np.shape(test_intensity_1[0]) == (5, 10)
assert np.all(np.isclose(np.sum(test_intensity_1[0], axis=0), np.array([47., 38., 56., 45., 54., 51., 46., 57., 50., 66.])))
assert np.all(np.isclose(np.sum(test_intensity_1[0], axis=1), np.array([98.,  96.,  96., 111., 109.])))
assert np.shape(test_intensity_1[1]) == (11,)
assert np.all(np.isclose(test_intensity_1[1], np.array([-5., -4., -3., -2., -1.,  0.,  1.,  2.,  3.,  4.,  5.])))
assert np.shape(test_intensity_1[2]) == (6,)
assert np.all(np.isclose(test_intensity_1[2], np.array([-2.5, -1.5, -0.5,  0.5,  1.5,  2.5])))
```

#### b) Wygenerowanie danych i wizualizacja

Wykorzystaj funkcje przygotowane w poprzednim podpunkcie zadania do zbadania intensywności wszystkich wczytanych do notatnika rozkładów punktowych. Ustaw podział obszaru na 40 podobszarów wzdłuż osi $X$ i 20 podobszarów wzdłuż osi $Y$. Analizę przeprowadź dla całego obszaru, na którym zostały wygenerowane punkty.

Przedstaw wyniki analizy graficznie w postaci histogramów dwuwymiarowych z nałożonymi na nie rozkładami punktów. Zestaw wyniki na pojedynczej figurze (siatka wykresów 2x2).
Dla wszystkich histogramów ustaw tę samą skalę kolorów z identycznym zakresem wartości ustalonym na podstawie minimalnej i maksymalnej wartości intensywności wyznaczonej dla kompletu analizowanych rozkładów.

<!-- grade_id: cell-046b66442d9f5ae1, points: 7, grade: true, solution: true -->
```python
# Create a list to store the DataFrames
df_list = [points_HP, points_UP, points_M, points_T]

# Define titles for the subplots
subplot_titles = ["points_HP", "points_UP", "points_M", "points_T"]

# Initialize variables to store the minimum and maximum values
x_min_all = float('inf')
x_max_all = float('-inf')
y_min_all = float('inf')
y_max_all = float('-inf')

# Calculate the overall minimum and maximum values across all datasets
for df in df_list:
    x_min = df['X'].min()
    x_max = df['X'].max()
    y_min = df['Y'].min()
    y_max = df['Y'].max()

    x_min_all = min(x_min_all, x_min)
    x_max_all = max(x_max_all, x_max)
    y_min_all = min(y_min_all, y_min)
    y_max_all = max(y_max_all, y_max)

# Call intensity_on_subregions for each dataset and collect results
intensity_data_list = []
for df in df_list:
    intensity_data = intensity_on_subregions(df, [40, 20], [x_min_all, x_max_all], [y_min_all, y_max_all])
    intensity_data_list.append(intensity_data)

# Determine common color scaling based on min and max intensities across all datasets
min_intensity = min(np.min(intensity_data[0]) for intensity_data in intensity_data_list)
max_intensity = max(np.max(intensity_data[0]) for intensity_data in intensity_data_list)

# Create a 2x2 grid of subplots with white background
fig, axes = plt.subplots(2, 2, figsize=(12, 8))

# Set the background color to white
fig.patch.set_facecolor('white')

# Plot histograms and overlay scatter plots for each dataset
for i, ax in enumerate(axes.flat):
    im = ax.imshow(intensity_data_list[i][0], cmap='viridis', origin='lower', extent=(x_min_all, x_max_all, y_min_all, y_max_all), vmin=min_intensity, vmax=max_intensity)
    ax.set_title(subplot_titles[i])
    ax.set_xlabel('X Coordinate')
    ax.set_ylabel('Y Coordinate')
    fig.colorbar(im, ax=ax, label='Intensity')

    # Overlay scatter plots of points on top of the histograms
    ax.scatter(df_list[i]['X'], df_list[i]['Y'], c='black', s=1, marker='o', label='Points')
    ax.legend()

# Add a common figure title
fig.suptitle('Point Distributions and Intensity', fontsize=14, fontweight='bold')

# Adjust subplots layout and spacing
plt.tight_layout(rect=[0, 0, 1, 0.95])

plt.show()
```

### Zadanie 2: Badanie intensywności procesu punktowego metodą funkcji jądrowych (20 pkt)

Przygotuj funkcję `intensity_on_kde()`, która będzie obliczać intensywność procesu punktowego z wykorzystaniem jądrowego estymatora intensywności.

Algorytm postępowania:
1) Generujemy punkty pomiarowe na regularnej siatce tak, że poszczególne punkty są oddalone od siebie o odległość $d_x$ w poziomie i $d_y$ w pionie, a skrajne punkty leżą na granicach analizowanego obszaru.
2) Dla każdego z wygenerowanych punktów pomiarowych szacujemy intensywność procesu za pomocą jądrowego estymatora intensywności danego wzorem:
$$ 	\grave{\lambda}(x,y) = \sum_{d_i < \tau} \frac{3}{\pi \tau^2}(1-\frac{d_i^2}{\tau^2})^2$$ 
gdzie: $d_i$ - odległość punktu od centrum funkcji jądrowej, $\tau$ - promień funkcji jądrowej.

Uwagi do wykonania zadania:
- Podczas pisania funkcji należy wykorzystać maksymalnie jedną pętlę `for` iterującą po kolejnych pozycjach funkcji jądrowej. Nie należy korzystać z pętli `for` podczas wyliczania wartości funkcji jądrowej w poszczególnych jej lokalizacjach.
- Funkcja przygotowana według powyższych wytycznych nie ma wbudowanego mechanizmu korekty na efekt brzegowy, w związku z czym będzie zaniżać wartość intensywności przy krawędziach analizowanego obszaru.

#### a) Przygotowanie funkcji

<!-- grade_id: cell-b00217faae7981e3, points: 0, grade: false, solution: true -->
```python
def intensity_on_kde(points, kernel_radius, grid, x_lim, y_lim):
    """
    Parameters
    -------
    points: DataFrame
        Tablica zawierająca dwie kolumny ze współrzędnymi punktów opisane jako "X" i "Y".
    kernel_radius: float
        Liczba dodatnia określająca promień funkcji jądrowej.
    grid: list
        Lista określająca liczbę punktów w poziomie i pionie, dla których wyliczane będą wartości intensywności.
        Przykład: [10, 10]
    x_lim: list
        Lista określająca zakres wartości współrzędnej X.
        Przykład: [0, 10]
    y_lim: list
        Lista określająca zakres wartości współrzędnej Y.
        Przykład: [0, 10]   

    Returns
    -------
    intensity_data: DataFrame
        Tablica zawierająca trzy kolumny - dwie ze współrzędnymi punktów, dla których wyliczone zostały wartości intensywności
        opisane jako "X" i "Y" oraz kolumnę z wartościami intensywności wyliczonymi dla tych współrzędnych opisaną jako "I".
    """
    
    '''
    jesli mamy jakis rozklad ktory bedziemy badac nakladamy na niego rowno rozlozony grid tylko pkt nie sa oddalone od siatki(?)
    centrujemy estymate jednorodna, w oparciu w b wykres konturowy z wypelnieniem 
    wygenerowane centrami liczymy odelgosci do wczytanych
    '''
   
    # YOUR CODE HERE
    
    # calculate all values of x and y coordinates
    x_coords = np.linspace( x_lim[0], x_lim[1], grid[0] )
    y_coords = np.linspace( y_lim[0], y_lim[1], grid[1] )
    x_matrix, y_matrix = np.meshgrid( x_coords, y_coords )
    intensity_data = pd.DataFrame( { "X": x_matrix.ravel(), "Y" : y_matrix.ravel() } ) 

    I_values= [] 
    
    for row in intensity_data.itertuples(index = False):
        
        # calculate distance between our centrum and input points 
        di = np.linalg.norm( points - row, axis = 1 )
        di = di[ di < kernel_radius]
        
        # estimate intensity of the process 
        I_values.append( np.sum( ( 3 / ( np.pi * kernel_radius ** 2 ) ) * ( 1 - di ** 2 / kernel_radius ** 2 ) ** 2 ) )
     
    intensity_data['I'] = I_values
    return intensity_data
```

<!-- grade_id: cell-fa0dc6e1c42a1ab1, points: 13, grade: true, solution: false -->
```python
# Komórka testowa
test_intensity_2 = intensity_on_kde(points_HP, kernel_radius=1.25, grid=[50, 25], x_lim=[-5, 5], y_lim=[-2.5, 2.5])
assert type(test_intensity_2) == pd.DataFrame
assert test_intensity_2.shape == (1250, 3)
assert list(test_intensity_2.columns) == ["X", "Y", "I"]
assert np.isclose(np.min(test_intensity_2["X"]), -5)
assert np.isclose(np.max(test_intensity_2["X"]), 5)
assert np.isclose(np.min(test_intensity_2["Y"]), -2.5)
assert np.isclose(np.max(test_intensity_2["Y"]), 2.5)
assert np.isclose(np.min(test_intensity_2["I"]), 4.063039686796769)
assert np.isclose(np.max(test_intensity_2["I"]), 15.939182630365504)
assert np.isclose(np.mean(test_intensity_2["I"]), 10.041689189726556)
```

#### b) Wygenerowanie danych i wizualizacja
Wykorzystaj przygotowaną funkcje do zbadania intensywności wszystkich wczytanych do notatnika rozkładów punktowych. Wykonaj obliczania dla promienia funkcji jądrowej równego 1.5 i siatki 200x100 punktów pomiarowych. Analizę przeprowadź dla całego obszaru, na którym zostały wygenerowane punkty.

Przedstaw wyniki analizy graficznie w postaci wykresów konturowych z wypełnieniem z nałożonymi na nie rozkładami punktów. Zestaw wyniki na pojedynczej figurze (siatka wykresów 2x2). Dla wszystkich histogramów ustaw tę samą skalę kolorów z identycznym zakresem wartości ustalonym na podstawie minimalnej i maksymalnej wartości intensywności wyznaczonej dla kompletu analizowanych rozkładów.

<!-- grade_id: cell-f8cd20d253f13d11, points: 7, grade: true, solution: true -->
```python
# YOUR CODE HERE
grid = [200, 100]
# Create a list to store the DataFrames
df_list = [points_HP, points_UP, points_M, points_T]

# Define titles for the subplots
subplot_titles = ["points_HP", "points_UP", "points_M", "points_T"]

# Initialize variables to store the minimum and maximum values
x_min_all = float('inf')
x_max_all = float('-inf')
y_min_all = float('inf')
y_max_all = float('-inf')

# Calculate the overall minimum and maximum values across all datasets
for df in df_list:
    x_min = df['X'].min()
    x_max = df['X'].max()
    y_min = df['Y'].min()
    y_max = df['Y'].max()

    x_min_all = min(x_min_all, x_min)
    x_max_all = max(x_max_all, x_max)
    y_min_all = min(y_min_all, y_min)
    y_max_all = max(y_max_all, y_max)

# Call intensity_on_subregions for each dataset and collect results
intensity_data_list = []
for df in df_list:
    intensity_data = intensity_on_kde(df, kernel_radius=1.5, grid=grid, x_lim=[x_min_all, x_max_all], y_lim=[y_min_all, y_max_all])
    intensity_data_list.append(intensity_data)

# Determine common color scaling based on min and max intensities across all datasets
min_intensity = min(np.min(intensity_data["I"]) for intensity_data in intensity_data_list)
max_intensity = max(np.max(intensity_data["I"]) for intensity_data in intensity_data_list)

# Create a 2x2 grid of subplots with white background
fig, axes = plt.subplots(2, 2, figsize=(12, 8))

# Set the background color to white
fig.patch.set_facecolor('white')

# Plot histograms and overlay scatter plots for each dataset
for i, ax in enumerate(axes.flat):
    intensity = intensity_data_list[i]["I"].values.reshape(grid[1], grid[0])
    im = ax.contourf(intensity, cmap='viridis', levels = 30, extent=(x_min_all, x_max_all, y_min_all, y_max_all), vmin=min_intensity, vmax=max_intensity)
    ax.set_title(subplot_titles[i])
    ax.set_xlabel('X Coordinate')
    ax.set_ylabel('Y Coordinate')
    fig.colorbar(im, ax=ax, label='Intensity')

    # Overlay scatter plots of points on top of the histograms
    ax.scatter(df_list[i]['X'], df_list[i]['Y'], c='black', s=1, marker='o', label='Points')
    ax.legend()

# Add a common figure title
fig.suptitle('Kernel Density Estimation (KDE) Intensity', fontsize=14, fontweight='bold')

# Adjust subplots layout and spacing
plt.tight_layout(rect=[0, 0, 1, 0.95])

plt.show()
```